<!-- grade_id: cell-7fdfde574a2fe74e, points: 0, grade: false, solution: true -->
```python
def add(a, b):
    """
    计算两个数字的和
    """
    ### BEGIN SOLUTION
    return a + b
    ### END SOLUTION

```

<!-- grade_id: cell-ca921c3516228b26, points: 10, grade: true, solution: false -->
```python
# 自动评分测试用例
assert add(1, 2) == 3, "基础正数测试失败"
assert add(-1, 1) == 0, "负数测试失败"
assert add(0, 0) == 0, "零测试失败"

### BEGIN HIDDEN TESTS
# 这里的测试用例对学生是隐藏的，用来防止学生硬编码作弊
assert add(100, 200) == 300
assert add(999, 1) == 1000
### END HIDDEN TESTS

print("本题所有测试通过！")

```